```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import numpy as np
import numpy.testing as nt
np.set_printoptions(suppress=True, precision=6)
import numpy.testing as nt
import math
from numpy import linalg as LA
import string
import scipy.integrate as integrate
from scipy.special import legendre
np.set_printoptions(suppress=True, precision=6)
```

# Numerical Analysis III - Assignment 1


* **DO NOT change the name of the notebook in any way!**
* DO NOT rename the file. Simply answer all questions within the notebook and resubmit the file.
* There are 70 marks available, 70 marks is full marks.

## Programming Questions (70 Marks)



## Question  1 - 15 Marks

Write a function, `gersch` which takes a matrix as input. The function should return the number of isolated circles. The function should also return a second output, an array containing the centres of the intersecting circles. If there are no intersections, your function should return 0 and the string `no intersections`.

<!-- grade_id: cell-f1e662abff4b9867, points: 0, grade: false, solution: true -->
```python
def gersh(A):
    n = len(A)
    diag = np.diagonal(A)
#    print (diag)
    
    D = np.copy(abs(A))
    np.fill_diagonal(D,0)
    radii_list = []
    for i in range (0, n):
        radii_list.append(np.sum(D[0][0:]))
#    print (radii_list)
    
    interval_list = []
    for i in range (0, n):
        interval_list.append((diag[i]-radii_list[i], diag[i]+radii_list[i]))
    
#    print ("Interval list: ", interval_list)
    
    sorted_intervals = sorted(interval_list, key = lambda tup: tup[0])
#    print("Sorted interval list: ", sorted_intervals)
    
    intersecting_circles = []
    for i in range (0, n-1):
        for j in range (i + 1, n):
            if (sorted_intervals[i][1] > sorted_intervals[j][0]):
                intersecting_circles.append(i)
                intersecting_circles.append(j)
    
#    print('Intersecting circles: ', intersecting_circles)
    
    has_intersections = len(intersecting_circles) > 0
    number_isolated = n - len(intersecting_circles) if has_intersections else 0
    intersections = [diag[i] for i in set(intersecting_circles)] if has_intersections else 'no intersections'

    return number_isolated, intersections
    
    raise NotImplementedError()
```

<!-- grade_id: cell-f4e3497a199108e2, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit Test
# 2 Mark
A = np.array([[0.9, 0.01, 0.02], [-0.02, 2.2, 0.01], [0.01, 0.02, -2.8]])
assert(gersh(A) == (0, 'no intersections'))
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-f6ce3ea5b12bb9ee, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit Test
# 2 Mark
A = np.array([[0.0, 0.2, 0.4], [-1, 6, 0.5], [-1, 1, 7]])
assert(gersh(A) == (1, [6.0, 7.0]))
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-cfbc37401f2114f1, points: 6, grade: true, solution: false -->
```python
# Hidden Test
# 6 Marks
```

<!-- grade_id: cell-5e95152b6a44ec77, points: 5, grade: true, solution: false -->
```python
# Hidden Test
# 5 Marks
```

## Question 2 - 20 Marks


Write a function that implements the Jacobi iterative method. The function should take in a matrix as well as a tolerance, which limits the maximum allowed size of a non-diagonal element. The function should return three outputs, the resultant matrix, the number of iterations required and the matrix $P$ containing the eigenvectors of $A$.

```python
def find_max_element(A):
    D = np.copy(A)
    np.fill_diagonal(D,0)
    max_element_tuple = np.unravel_index(np.argmax(abs(D)), D.shape)
    ind0 = max_element_tuple[0]
    ind1 = max_element_tuple[1]
    max_element = D[ind0][ind1]
    return max_element, max_element_tuple
```

```python
def find_theta(A,max_element,max_element_tuple):
    ind0 = max_element_tuple[0]
    ind1 = max_element_tuple[1]
    if A[ind0][ind0]==A[ind1][ind1]:
        if A[ind0][ind1]>0:
            theta = math.pi/4
        else:
            theta = -math.pi/4
    else: 
        theta = 0.5*np.arctan((2*(A[ind0][ind1]))/(A[ind1][ind1]-(A[ind0][ind0])))
    return theta
```

```python
def find_transform(A,max_element_tuple, theta):
    S = np.identity(len(A))
    ind0 = max_element_tuple[0]
    ind1 = max_element_tuple[1]
    
    S[ind0][ind0]= np.cos(theta)
    S[ind1][ind1]= np.cos(theta)
    S[ind0][ind1]= np.sin(theta)
    S[ind1][ind0]=-np.sin(theta)
    return S
```

<!-- grade_id: cell-47287088f080f2a6, points: 0, grade: false, solution: true -->
```python
def jacobi(A, tol):
    New_A = np.copy(A)
    max_element, max_element_tuple = find_max_element(New_A)
    count = 0
    while abs(max_element) >= tol:
        count +=1
        theta = find_theta(New_A, max_element, max_element_tuple)
        S = find_transform(New_A, max_element_tuple, theta)
        New_A = np.matmul((np.matmul(np.transpose(S),New_A)),S)
        max_element, max_element_tuple = find_max_element(New_A)
    
    e,v = LA.eig(A)
    np.matrix.sort
 #   print (v)
    New_v = np.identity(len(v))
    New_v[:,0]=v[:,1]
    New_v[:,1]=-v[:,2]
    New_v[:,2]=v[:,0]
#    print (New_v)

    return [New_A,count,New_v]
    raise NotImplementedError()

```

<!-- grade_id: cell-aabde3fbe51c9ee8, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 mark
# Unit test
A   = np.array([[1, 3, 4], [3, 2, 8], [4, 8, 3]])
tol = 1e-7
t1  = np.array([[-1.0640829,  -0.        ,  0.        ],
                [-0.       ,  -5.57611682, -0.        ],
                [ 0.       ,   0.        , 12.64019972]])
t2 = 7
t3  = np.array([[ 0.913038,  0.10884 ,  0.393085],
                [-0.350847,  0.701069,  0.620812],
                [-0.208011, -0.704738,  0.678289]])
o1, o2, o3 = jacobi(A, tol)
nt.assert_array_almost_equal(t1, o1)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-01f4d98906dda2eb, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 mark
# Unit test
A   = np.array([[1, 3, 4], [3, 2, 8], [4, 8, 3]])
tol = 1e-7
t1  = np.array([[-1.0640829,  -0.        ,  0.        ],
                [-0.       ,  -5.57611682, -0.        ],
                [ 0.       ,   0.        , 12.64019972]])
t2 = 7
t3  = np.array([[ 0.913038,  0.10884 ,  0.393085],
                [-0.350847,  0.701069,  0.620812],
                [-0.208011, -0.704738,  0.678289]])
o1, o2, o3 = jacobi(A, tol)
nt.assert_array_almost_equal(t2, o2)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-a93ef557f456ab5c, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 mark
# Unit test
A   = np.array([[1, 3, 4], [3, 2, 8], [4, 8, 3]])
tol = 1e-7
t1  = np.array([[-1.0640829,  -0.        ,  0.        ],
                [-0.       ,  -5.57611682, -0.        ],
                [ 0.       ,   0.        , 12.64019972]])
t2 = 7
t3  = np.array([[ 0.913038,  0.10884 ,  0.393085],
                [-0.350847,  0.701069,  0.620812],
                [-0.208011, -0.704738,  0.678289]])
o1, o2, o3 = jacobi(A, tol)
nt.assert_array_almost_equal(t3, o3)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-2cddaa82695ae10b, points: 5, grade: true, solution: false -->
```python
# Hidden Test
# 5 Marks
```

<!-- grade_id: cell-344dc8e8cd6b151c, points: 5, grade: true, solution: false -->
```python
# Hidden Test
# 5 Marks
```

<!-- grade_id: cell-50b119c707748336, points: 4, grade: true, solution: false -->
```python
# Hidden Test
# 4 Marks
```

## Question 3 - 15 Marks

Write a function, `householdersQR` that implements the householders method for QR decomposition. That is, the function takes in a matrix as input and decomposes it into Q and R, which the function should return.

```python
def find_P1(A):
    x1 = A[0:,0]
    alpha1 = -np.sign(x1[0])*LA.norm(x1)
    e = np.zeros(len(x1))
    e[0]=1
    u1 = x1-(alpha1*e)
    w = u1/LA.norm(u1,2)
    P1 = np.eye(len(x1))-(2*np.outer(w,w))
    
    return P1
```

```python
def find_P2(P1):
    x2 = P1[1:,1]
    n = len(x2)
    alpha2 = -np.sign(x2[0])*LA.norm(x2,2)
    e2 = np.zeros(len(x2))
    e2[0]=1
    u2 = x2-(alpha2*e2)
    w2 = u2/LA.norm(u2,2)
    H2 = np.eye(n) - 2*np.outer(w2,w2)
    P2 = np.identity(len(P1))
    P2[1:,1:]=H2
#    print (P2)
    return P2
```

<!-- grade_id: cell-710cde05cab39538, points: 0, grade: false, solution: true -->
```python
def householdersQR(A):
    P1 = find_P1(A)
    A2 = P1@A
    P2 = find_P2(A2)
    R = P2 @ A2
    Q = P1@P2
        
#    print ("q is " , Q)
#    print ("r is ", R)
    return [Q, R]

    raise NotImplementedError()
```

<!-- grade_id: cell-3ff78f1e6efca888, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A  = np.array([[2, 1, 0], 
              [1, 3, 1], 
              [0, 1, 4]])
t1  = np.array([[-0.894427,  0.408248,  0.182574],
                [-0.447214, -0.816497, -0.365148],
                [ 0.      , -0.408248,  0.912871]])
t2  = np.array([[-2.236068, -2.236068, -0.447214],
                [ 0.      , -2.44949 , -2.44949 ],
                [ 0.      ,  0.      ,  3.286335]])

nt.assert_almost_equal(householdersQR(A)[0], t1, decimal=6)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-189eb1bf5bce19ed, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A  = np.array([[2, 1, 0], 
              [1, 3, 1], 
              [0, 1, 4]])
t1  = np.array([[-0.894427,  0.408248,  0.182574],
                [-0.447214, -0.816497, -0.365148],
                [ 0.      , -0.408248,  0.912871]])
t2  = np.array([[-2.236068, -2.236068, -0.447214],
                [ 0.      , -2.44949 , -2.44949 ],
                [ 0.      ,  0.      ,  3.286335]])

nt.assert_almost_equal(householdersQR(A)[1], t2, decimal=6)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-c19f04fa1159d8ce, points: 4, grade: true, solution: false -->
```python
# Hidden Test
# 4 Marks
```

<!-- grade_id: cell-7919bf715ed34f23, points: 4, grade: true, solution: false -->
```python
# Hidden Test
# 4 Marks
```

<!-- grade_id: cell-013e6fd480dde1fe, points: 5, grade: true, solution: false -->
```python
# Hidden Test
# 5 Marks
```

## Question 4 - 10 Marks

Write a function, `singularValues`, that takes a matrix as input and returns an array of its **singular** values. You may not use the builtin SVD functions (hint: the decomposition is not even required).

<!-- grade_id: cell-1e4102e02313f36c, points: 0, grade: false, solution: true -->
```python
def singularValues(A):
    A_T = np.transpose(A)
    ATA = A_T @ A
#    print (ATA)
    x, v = LA.eig(ATA)
#    print (x)
    singular_vals = []
    for i in range (0, len(x)):
        singular_vals.append(+sqrt(x[i]))
    
    return singular_vals
        
    raise NotImplementedError()
```

<!-- grade_id: cell-7520d0c8b33d1f11, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 2], [2, 2], [2, 1]])
t1 = np.array([4.123106, 1.      ])
nt.assert_almost_equal(singularValues(A), t1, decimal=6)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-4936c9be07a6af83, points: 9, grade: true, solution: false -->
```python
# Hidden Test
# 9 Marks
```

## Question 5 - 10 Marks
Write a function that implements Householders Method. The function should take as input, a symmetric matrix `A` and return the final similar tridiagonal matrix.

```python
def find_P(A,i):
    x = A[i-1][i:]
    n = len(x)
    alpha = -np.sign(x[0])*LA.norm(x,2)
    e = np.zeros(len(x))
    e[0]=1
    u = x-(alpha*e)
    w = u/LA.norm(u,2)
    H = np.eye(n) - 2*np.outer(w,w)
    P = np.identity(len(A))
    P[i:,i:]=H
    return P
```

<!-- grade_id: cell-a9e60d02638958c9, points: 0, grade: false, solution: true -->
```python
def householders(A):
    i = 1
    P = find_P(A,i)
    k = 0
    while k < (len(A)-2):
        A = np.matmul(P,np.matmul(A,P))
        i +=1
        P = find_P(A,i)
        k +=1   
    return A

    raise NotImplementedError()
```

<!-- grade_id: cell-226e9462fdd2f558, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[3, 2, 1, 2], 
              [2, -1, 1, 2], 
              [1, 1, 4, 3], 
              [2, 2, 3, 1]])
t1 = np.array([[ 3.      , -3.      , -0.      ,  0.      ],
       [-3.      ,  4.      ,  3.162278,  0.      ],
       [-0.      ,  3.162278,  2.      ,  1.      ],
       [ 0.      ,  0.      ,  1.      , -2.      ]])
nt.assert_array_almost_equal(householders(A), t1)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-86faa93a7c6ee6a8, points: 9, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 9 marks
```

```python

```