```python
# Initialisation Cell
from matplotlib import pyplot as plt
import math
from math import *
import numpy as np
from numpy import linalg as LA
import numpy.testing as nt
```

# Numerical Analysis III - Lab 1


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.15.4/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

# Warm-up Exercises

Complete the following warm-up tasks without the use of numpy.

## Question 1

Given an array of integers, write a function to compute the sum of its elements. You may not use the `sum` function.

<!-- grade_id: cell-dd4f4abacf2c1358, points: 0, grade: false, solution: true -->
```python
# Question 1
def simpleArraySum(ar):
    sumup=0
    length = len(ar)
    for num in ar:
        sumup +=num
    
    return sumup
        
    raise NotImplementedError()
```

<!-- grade_id: cell-57df763d59cc738e, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test = [1, 2, 3 , 4, 10, 11]
assert(simpleArraySum(test) == 31)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-5609a53f91ced436, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Question 2

Write a function which takes in a matrix and computes the absolute difference between the sums of its diagonals (i.e. left-to-right and right-to-left).

<!-- grade_id: cell-9216b1ad99ff18e8, points: 0, grade: false, solution: true -->
```python
# Question 2

def diagonalDifference(arr):
    sum1 = np.trace(arr)  
#    print (sum1)
    newArr = np.flip(arr,axis=1)
#    print (newArr)
    sum2 = np.trace(newArr) 
    diff = abs(sum2-sum1)
    return diff
    raise NotImplementedError()
    
```

<!-- grade_id: cell-673842ebeecc4f16, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test =  [[11, 2,   4],
         [4,  5,   6],
         [10, 8, -12]]

assert(diagonalDifference(test) == 15)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-ee3f9ffebbaed8ae, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Question 3

Write a function that takes as input an array of integers. The function should compute the fractions of its elements that are positive, negative and are zeros. These fractions should be returned rounded to 4 decimal places.

<!-- grade_id: cell-54809b688e68a9f0, points: 0, grade: false, solution: true -->
```python
# Question 3

def plusMinus(arr):
    length = len(arr)
    count1 = 0
    count2 = 0
    count3 = 0
    for num in (arr):
        if num >0:
            count1+=1
        elif num<0:
            count2 +=1
        else:
            count3+=1
    
    return [round(count1/length,4), round(count2/length,4), round(count3/length,4)]
    raise NotImplementedError()
```

<!-- grade_id: cell-ffc6e1a3d4e7eb1f, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
arr = [-4, 3, -9, 0, 4, 1]
nt.assert_array_almost_equal(plusMinus(arr),(0.5, 0.3333, 0.1667))
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-e1993ec3793d8e84, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Main Exercises

### Question 1

Write a function that takes as input a matrix and determines whether or not the matrix is symmetric. The function should return the strings `not symmetric` or `symmetric`. You may not use any obvious builtin function.


<!-- grade_id: cell-bb79f3b8b5237325, points: 0, grade: false, solution: true -->
```python
# Question 1
def symCheck(A):
    
#    for i in range(A): 
#    for j in range(A): 
#        tr[i][j] = A[j][i]
#     if (not square)
#         return "not symmetric"
    A_transpose = A.transpose()
    return "symmetric" if np.all(A == A_transpose) else "not symmetric"
    raise NotImplementedError()
    
```

<!-- grade_id: cell-6ceae1a45ca31ec3, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 2, 3], [2, 1, 3]])
assert(symCheck(A) == 'not symmetric')
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
/Users/charlottesavage/opt/anaconda3/lib/python3.7/site-packages/ipykernel_launcher.py:10: DeprecationWarning: elementwise comparison failed; this will raise an error in the future.
  # Remove the CWD from sys.path while we load stuff.
```

<!-- grade_id: cell-930f852483475e58, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 2, 3], [2, 1, 3],  [2, 1, 3]])
assert(symCheck(A) == 'not symmetric')
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-90a862f789352887, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 2, 3], [2, 1, 3],  [3, 3, 3]])
assert(symCheck(A) == 'symmetric')
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-f9e2f8e214568ddc, points: 3, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 3 marks
```

<!-- grade_id: cell-77a576b296165b9c, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

### Question 2
Write a function that takes in a matrix and determines whether or not the matrix is diagonally dominant. If the matrix is not strictly diagonally dominant, then the function should return `not diagonally dominant`, else the function must return `diagonally dominant`.

<!-- grade_id: cell-821c25ec75fcec98, points: 0, grade: false, solution: true -->
```python
# Question 2
def diagonallyDominant(A):
    diag = abs(np.diagonal(A))
    if A.shape[0]!=A.shape[1]:
        return ('not diagonally dominant')
    else:
        for i in range (0,len(diag)):
            if np.all(abs(sum(A[i])-diag[i]) <= diag[i]):
                return 'diagonally dominant'
            else:
                return "not diagonally dominant"

    raise NotImplementedError()  
```

<!-- grade_id: cell-f964ef3cb505deaa, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 2, 3], [2, 1, 3]])
diagonallyDominant(A)
assert(diagonallyDominant(A) == 'not diagonally dominant')
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-a14be8e4293e1c64, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[3, -2, 1], [1, -3, 2], [-1, 2, 4]])
assert(diagonallyDominant(A) == 'diagonally dominant')
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-c5f8c7cd87dbe271, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

<!-- grade_id: cell-34f4207fc3262f08, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

### Question 3

Write a function that implements the Power method for computing the largest Eigenvalue (recall from Numerics II). The function should take in a matrix A, an initial guess `x`, the number of iterations to perform `k` and return the eigenvalue and eigenvector pair (using the $\ell_2$ norm).

<!-- grade_id: cell-73cfeb4a00fb49f4, points: 0, grade: false, solution: true -->
```python
def power_iteration(A, x, k):
    for i in range (k):
        u = x/(LA.norm(x))
#        print (u)
        x = np.matmul(A,u)
#        print (x)
    lamb = LA.norm(x)
    u = x/(LA.norm(x))
    return [lamb,u]
    
    raise NotImplementedError()
```

<!-- grade_id: cell-d12def2ae18d05fe, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 3], [2, 2]])
x = np.array([-5, 5])
k = 10
nt.assert_approx_equal(power_iteration(A, x, k)[0], 3.999980924767464)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-94a6f86cc264baac, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 3], [2, 2]])
x = np.array([-5, 5])
k = 10
nt.assert_array_almost_equal(power_iteration(A, x, k)[1], np.array([0.70710341, 0.70711015]))
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-0cb5374bbbf85aa8, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

<!-- grade_id: cell-bb9f3b23195b1a8d, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

### Question 4

Write a function that implements the Jacobi iterative method which iterates while the magnitude of the largerst off diagonal is greater than some tolerance. The function should take in a matrix as well as a tolerance. The function should return the resultant matrix and the number of iterations required.

<!-- grade_id: cell-354a498585cf3634, points: 0, grade: false, solution: true -->
```python
def find_max_element(A):
    D = np.copy(A)
    np.fill_diagonal(D,0)
 #   print (D)
 #   abs_max_element = np.max(abs(D))
    max_element_tuple = np.unravel_index(np.argmax(abs(D)), D.shape)
 #   print (max_element_tuple)
    ind0 = max_element_tuple[0]
    ind1 = max_element_tuple[1]
    max_element = D[ind0][ind1]
#    print ("Max element is " + str(max_element))
    return max_element, max_element_tuple
```

```python
def find_theta(A,max_element,max_element_tuple):
    ind0 = max_element_tuple[0]
    ind1 = max_element_tuple[1]
    if A[ind0][ind0]==A[ind1][ind1]:
        if A[ind0][ind1]>0:
            theta = math.pi/4
        else:
            theta = -math.pi/4
    else: 
        theta = 0.5*np.arctan((2*(A[ind0][ind1]))/(A[ind1][ind1]-(A[ind0][ind0])))
    print ("this is theta "+str(theta))
    return theta
```

```python
def find_transform(A,max_element_tuple, theta):
    S = np.identity(len(A))
    ind0 = max_element_tuple[0]
    ind1 = max_element_tuple[1]
    
    S[ind0][ind0]= np.cos(theta)
    S[ind1][ind1]= np.cos(theta)
    S[ind0][ind1]= np.sin(theta)
    S[ind1][ind0]=-np.sin(theta)
#    print (S)
    return S
```

```python
def jacobi(A, tol):
    max_element, max_element_tuple = find_max_element(A)
    count = 0
    while abs(max_element) >= tol:
        count +=1
        theta = find_theta(A, max_element, max_element_tuple)
        S = find_transform(A, max_element_tuple, theta)
        A = np.matmul((np.matmul(np.transpose(S),A)),S)
        max_element, max_element_tuple = find_max_element(A)

    return [A,count]
    raise NotImplementedError()
```

<!-- grade_id: cell-960bb781516adac2, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A   = np.array([[1, 3, 4], [3, 2, 8], [4, 8, 3]])
tol = 1e-7
t1  = np.array([[-1.0640829,  -0.        ,  0.        ],
                [-0.       ,  -5.57611682, -0.        ],
                [ 0.       ,   0.        , 12.64019972]])
t2 = 7
o1, o2 = jacobi(A, tol)
nt.assert_array_almost_equal(t1, o1)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-0bb7f9c590308d45, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A   = np.array([[1, 3, 4], [3, 2, 8], [4, 8, 3]])
tol = 1e-7
t1  = np.array([[-1.0640829,  -0.        ,  0.        ],
                [-0.       ,  -5.57611682, -0.        ],
                [ 0.       ,   0.        , 12.64019972]])
t2 = 7
o1, o2 = jacobi(A, tol)
assert(abs(o2 - t2) <= 1)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-1c7903e40fb17c04, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

<!-- grade_id: cell-fce31b93e88cf1b9, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

```python

```