```python
# Initialisation Cell
from matplotlib import pyplot as plt
import math
from math import *
import numpy as np
from numpy import linalg as LA
import numpy.testing as nt
np.set_printoptions(suppress=True, precision=6)
```

# Numerical Analysis III - Lab 2


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.15.4/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

# Warm-up Exercises

Complete the following warm-up tasks without the use of numpy.

## Question 1

Write a function that takes as input a vector `arr`. The function should return both the minimum and maximum (in this order) possible sums utilising `n - 1` elements of the input vector, where `n` is the length of the vector.

<!-- grade_id: cell-7c407948d50a4b28, points: 0, grade: false, solution: true -->
```python
# Question 1
def miniMaxSum(arr):
    arr.sort()
    sum1 =0
    for num in arr:
        sum1 +=num
    return [sum1-arr[-1],sum1-arr[0]]
            
    raise NotImplementedError()
```

<!-- grade_id: cell-c3d4eb53c5473254, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
arr = [1, 2, 3, 4, 5]
assert(miniMaxSum(arr)[0] == 10)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-7ceb3dbc5c280a5b, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
arr = [1, 2, 3, 4, 5]
assert(miniMaxSum(arr)[1] == 14)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-9113e57a7387853a, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

<!-- grade_id: cell-e655ad8b569be729, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Question 2

Write a function that takes as input a string `s` containing a time in 12-hour format (i.e. hh:mm:ssAM or hh:mm:ssPM). The function should return the time as a string in 24-hour format.

<!-- grade_id: cell-e49c3e784dfa3bcf, points: 0, grade: false, solution: true -->
```python
# Question 2
def timeConversion(s):
    if str(s[-2:])=="PM"and (s[0:2]!='12'):
        return (str(int(s[0:2])+12)+str(s[2:-2]))    
    elif str(s[-2:])=="AM"and (s[0:2]=="12"):
        return ("00"+str(s[2:-2]))
    else:  
        return(s[0:-2])       
    raise NotImplementedError()
```

<!-- grade_id: cell-f686d2ee886a505b, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
s = '06:40:03AM'
assert(timeConversion(s) == '06:40:03')
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-8e6bc11a5ef23067, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
s = '12:45:54AM'
assert(timeConversion(s) == '00:45:54')
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-8106834aef02cfa5, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

<!-- grade_id: cell-7d12efcee6350eb5, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Question 3

Write a function that takes in an array, `ar`, of $n$ integers and a positive integer `k`. Find the number of $(i, j)$ pairs where $i < j$ and $ar(i) + ar(j)$ is divisible by `k`. The function should return the answer as a scalar.

<!-- grade_id: cell-9c5e35032c75619c, points: 0, grade: false, solution: true -->
```python
# Question 3
def divisibleSumPairs(k, ar):
    n = len(ar)
    count = 0
    for i in range (0, n):
        for j in range(i+1,n):
            if (ar[i]+ar[j])%k==0:
                count +=1
    return count
    raise NotImplementedError()
```

<!-- grade_id: cell-8b3a76ac686c4eb3, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
k    = 3
test = [1, 3, 2, 6, 1, 2]
assert(divisibleSumPairs(k, test) == 5)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-3e916b818b7d8520, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
k    = 3
test = [1, 3, 2, 6, 1, 2, 5, 9, 4]
assert(divisibleSumPairs(k, test) == 12)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-cb69bc00177f59d2, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

<!-- grade_id: cell-3855c1ecf8f4715e, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Main Exercises

### Question 1

Write a function that implements Givens Method. The function should take as input, a symmetric matrix `A` and return the final similar tridiagonal matrix.


```python
def annihilate_entry(A, i, j):
    i_prime = i + 1
    theta = -np.arctan(A[i][j] / A[i][i_prime])
#     print('Theta: ', theta, ', want: ', -math.pi / 4)
    P = np.identity(len(A))
#     print('Cos : ', np.cos(theta), 'sin: ', np.sin(theta))
    P[i_prime][i_prime] = np.cos(theta)
    P[j][j] = np.cos(theta)
    P[i_prime][j] = np.sin(theta)
    P[j][i_prime] = -P[i_prime][j]
#     print(P)
    A_prime = P.transpose() @ A @ P
    return A_prime
```

```python
def find_next_entry_to_kill(A):
    for i in range(len(A)):
        for j in range(i + 2, len(A)):
            if A[i][j] != 0:
                return i, j
    return -1, -1
```

<!-- grade_id: cell-73f72b49445c73cb, points: 0, grade: false, solution: true -->
```python
# Question 1
def givens(A):
    i, j = find_next_entry_to_kill(A)
    while i > -1:
        A = annihilate_entry(A, i, j)
#         print('A this round: ', A)
        i, j = find_next_entry_to_kill(A)
#         print('Next non zero : ', i, j)
    
    return A
    raise NotImplementedError()
```

<!-- grade_id: cell-53cd005a7b12cff9, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, sqrt(2), sqrt(2), 2], 
              [sqrt(2), -sqrt(2), -1, sqrt(2)], 
              [sqrt(2), -1, sqrt(2), sqrt(2)], 
              [2, sqrt(2), sqrt(2), -3]])
t1 = np.array([[ 1.      ,  2.828427,  0.      ,  0.      ],
               [ 2.828427,  0.      ,  1.414214, -0.      ],
               [ 0.      ,  1.414214, -0.5     ,  2.5     ],
               [ 0.      , -0.      ,  2.5     , -2.5     ]])
nt.assert_array_almost_equal(givens(A), t1)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-d1e7ed434cfed791, points: 9, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 9 marks
```

### Question 2
Write a function that implements Householders Method. The function should take as input, a symmetric matrix `A` and return the final similar tridiagonal matrix.

```python
def find_P(A,i):
    x = A[i-1][i:]
    n = len(x)
    alpha = -np.sign(x[0])*LA.norm(x,2)
    e = np.zeros(len(x))
    e[0]=1
    u = x-(alpha*e)
    w = u/LA.norm(u,2)
    H = np.eye(n) - 2*np.outer(w,w)
    P = np.identity(len(A))
    P[i:,i:]=H
    return P
```

<!-- grade_id: cell-87e89ffe6d3f1ae9, points: 0, grade: false, solution: true -->
```python
def householders(A):
    i = 1
    P = find_P(A,i)
    k = 0
    while k < (len(A)-2):
        A = np.matmul(P,np.matmul(A,P))
        i +=1
        P = find_P(A,i)
        k +=1   
    return A

    raise NotImplementedError()
```

<!-- grade_id: cell-853d580430377cf0, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[3, 2, 1, 2], 
              [2, -1, 1, 2], 
              [1, 1, 4, 3], 
              [2, 2, 3, 1]])
t1 = np.array([[ 3.      , -3.      , -0.      ,  0.      ],
       [-3.      ,  4.      ,  3.162278,  0.      ],
       [-0.      ,  3.162278,  2.      ,  1.      ],
       [ 0.      ,  0.      ,  1.      , -2.      ]])
nt.assert_array_almost_equal(householders(A), t1)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-6f854c4e52dccea2, points: 9, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 9 marks
```

## Question 3

Write a function that implements the QR method without shifting. The function should take as inputs, the matrix `A` and tolerance, `tol`. The function should return the approximate eigenvalues. It should also utilise the Householders method to do the required factorisation. Finally, the function should also return the number of iterations required until all off-diagonal elements of `A` are less than `tol`.

```python
def find_P1(A):
    x1 = A[0][0:]
    alpha1 = -np.sign(x1[0])*LA.norm(x1)
    e = np.zeros(len(x1))
    e[0]=1
    u1 = x1-(alpha1*e)
    w = u1/LA.norm(u1,2)
    P1 = np.eye(len(x1))-(2*np.outer(w,w))
    
    return P1
```

```python
def find_P2(P1):
    x2 = P1[1][1:]
    n = len(x2)
    alpha2 = -np.sign(x2[0])*LA.norm(x2,2)
    e2 = np.zeros(len(x2))
    e2[0]=1
    u2 = x2-(alpha2*e2)
    w2 = u2/LA.norm(u2,2)
    H2 = np.eye(n) - 2*np.outer(w2,w2)
    P2 = np.identity(len(P1))
    P2[1:,1:]=H2
#    print (P2)
    return P2
```

<!-- grade_id: cell-c74287fcc3b8c4fa, points: 0, grade: false, solution: true -->
```python
def QR(A, tol):
    count = 0
    while abs(A[0][1]) > tol:
        P1 = find_P1(A)
        A2 = np.matmul(P1,A)
        P2 = find_P2(A2)
        R = P2 @ P1 @ A
        Q = np.matmul(P1, P2)
        
        A = np.matmul(R,Q)
        count +=1
        diag = np.diagonal(A)
        
    return [diag, count]
    
    raise NotImplementedError()
```

<!-- grade_id: cell-311e18a74770bde8, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A   = np.array([[2, 1, 0], [1, 3, 1], [0, 1, 4]])
tol = 1e-3
t1  = np.array([4.73205 , 3.      , 1.267949])
t2  = 19
nt.assert_array_almost_equal(QR(A, tol)[0], t1)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-21d8ebcd70cbf461, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A   = np.array([[2, 1, 0], 
                [1, 3, 1], 
                [0, 1, 4]])
tol = 1e-3
t1  = np.array([4.73205 , 3.      , 1.267949])
t2  = 19
assert(abs(QR(A, tol)[1] - t2) <= 2)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-09193f4d5c13da86, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

<!-- grade_id: cell-8aef83b083fb306b, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Question 4

Redo question (3), however, now implement the QR method with shifting using method 1 from the notes.

```python
def find_P1(A):
    x1 = A[0][0:]
    alpha1 = -np.sign(x1[0])*LA.norm(x1)
    e = np.zeros(len(x1))
    e[0]=1
    u1 = x1-(alpha1*e)
    w = u1/LA.norm(u1,2)
    P1 = np.eye(len(x1))-(2*np.outer(w,w))
    
    return P1
```

```python
def find_P2(P1):
    x2 = P1[1][1:]
    n = len(x2)
    alpha2 = -np.sign(x2[0])*LA.norm(x2,2)
    e2 = np.zeros(len(x2))
    e2[0]=1
    u2 = x2-(alpha2*e2)
    w2 = u2/LA.norm(u2,2)
    H2 = np.eye(n) - 2*np.outer(w2,w2)
    P2 = np.identity(len(P1))
    P2[1:,1:]=H2
#    print (P2)
    return P2
```

<!-- grade_id: cell-4b9df052ba2eba58, points: 0, grade: false, solution: true -->
```python
def QRshifting(A, tol):
    count = 0
    while abs(A[0][1]) > tol:
        n = len(A)
        c = A[n-1][n-1]
        cI = c*np.identity(n)
        A1 = A-cI
        
        P1 = find_P1(A1)

        Aprime = np.matmul(P1,A1)
        P2 = find_P2(Aprime)
        R = P2 @ P1 @ A1
        Q = np.matmul(P1, P2)

        A = cI+(np.matmul(R,Q))
        count +=1
        diag = np.diagonal(A)
#    print (A)
#    print (count)
        
    return [diag, count]
    raise NotImplementedError()

```

<!-- grade_id: cell-e492a96886ad818b, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A   = np.array([[2, 1, 0], [1, 3, 1], [0, 1, 4]])
tol = tol
t1  = np.array([1.26795 , 3.      , 4.732051])
t2  = 10
nt.assert_array_almost_equal(QRshifting(A, tol)[0], t1)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-4375df249928850b, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A   = np.array([[2, 1, 0], [1, 3, 1], [0, 1, 4]])
tol = tol
t1  = np.array([1.26795 , 3.      , 4.732051])
t2  = 10
assert(abs(QRshifting(A, tol)[1] - t2) <= 2)
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-d1a3f937ef8838a7, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

<!-- grade_id: cell-febac481f10aa638, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

```python

```