```python
# Initialisation Cell
from matplotlib import pyplot as plt
from math import *
import numpy as np
import numpy.testing as nt
np.set_printoptions(suppress=True, precision=6)
```

# Numerical Analysis III - Lab 3


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.15.4/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

# Warm-up Exercises

Complete the following warm-up tasks without the use of numpy.

## Question 1

Write a function that takes in a binary array of length `n`. The function must compute the minimum number of steps to navigate the array, where zero elements are values that may be traversed and ones **taboo** and must be avoided. The maximum size of the step allowable is 2, i.e. a shift of 1 or 2. For example, if $c = [0, 1, 0, 0, 0, 1, 0]$ indexed from $0, \ldots, 6$. The possible paths are $0 \rightarrow 2 \rightarrow 4 \rightarrow 6$ or $0 \rightarrow 2 \rightarrow 3 \rightarrow 4 \rightarrow 6$. The first path requires 3 shifts and the second 4 shifts. Therefore, the function should return the shortest number of shifts.

<!-- grade_id: cell-e90753999f991276, points: 0, grade: false, solution: true -->
```python
# Question 1
def jumping(c):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-461a79a8dfe5b687, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
c = [0, 1, 0, 0, 0, 1, 0]
assert(jumping(c) == 3)
print('Test case passed!!!')
```

<!-- grade_id: cell-88e41979df52f5fb, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Question 2

Write a function that takes in two arrays and determines all integers satisfying the following:
* The elements of the first array that are all factors of the integer being considered.
* The integer being considered is a factor of all elements in the second array. 
    
Your function must the amount of such numbers. For example, if $a = [2, 6]$ and $b = [24, 36]$, there are 2 numbers between them: 6 and 12, i.e. 6 % 2 = 0, 6 % 6 = 0, 24 % 6 = 0 and 36 % 6 for the first value. Similarly, 12 % 2 = 0, 12 % 6 = 0,  24 % 12 = 0 and 36 % 12 = 0.

<!-- grade_id: cell-d952af01780c2829, points: 0, grade: false, solution: true -->
```python
# Question 2
def getTotalX(a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e94fea2c36d576b8, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
a = [2, 6]
b = [24, 36]
assert(getTotalX(a, b) == 2)
print('Test case passed!!!')
```

<!-- grade_id: cell-cc0db9a5d2599d4c, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Question 3

Write a function that takes as inputs, a word (as a string) and an array of letter heights. The function must compute the area required to highlight the entire word. For example, if the string is `'abc'`, and the specific heights of `a`, `b` and `c` are 1, 3, and 1 mm respectively - then the background area to highlight the entire string is the length of the string times the maximum height, therefore $3 \times 3 = 9$mm.

<!-- grade_id: cell-d01708c3941de3e3, points: 0, grade: false, solution: true -->
```python
# Question 3

def highlighter(h, word):
    # YOUR CODE HERE
    raise NotImplementedError()


```

<!-- grade_id: cell-e6e9bbd510e10288, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test = 'abc'
h    = [1, 3, 1, 3, 1, 4, 1, 3, 2, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 7]
assert(highlighter(h, test) == 9)
print('Test case passed!!!')
```

<!-- grade_id: cell-aab8ed33c4131a5c, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Main Exercises

Complete the following problems:

## Question 1

Write a function that computes the chebyshev polynomial of degree `n` on the vector domain `x`. The function should return the $y(x)$ corresponding to the vector input `x`.

<!-- grade_id: cell-dae3990afd85b835, points: 0, grade: false, solution: true -->
```python
def chebyshev(n, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b6c3f19eb0fec04f, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
n = 3
x = np.arange(-1, 1.05, 0.05)
t1 = np.array([-1.    , -0.5795, -0.216 ,  0.0935,  0.352 ,  0.5625,  0.728 ,
            0.8515,  0.936 ,  0.9845,  1.    ,  0.9855,  0.944 ,  0.8785,
            0.792 ,  0.6875,  0.568 ,  0.4365,  0.296 ,  0.1495, -0.    ,
           -0.1495, -0.296 , -0.4365, -0.568 , -0.6875, -0.792 , -0.8785,
           -0.944 , -0.9855, -1.    , -0.9845, -0.936 , -0.8515, -0.728 ,
           -0.5625, -0.352 , -0.0935,  0.216 ,  0.5795,  1.    ])
nt.assert_almost_equal(t1, chebyshev(n, x))
print('Test case passed!!!')
```

<!-- grade_id: cell-cde2778033df39d4, points: 9, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 9 marks
```

## Question 2

Write a function which approximates functions using chebyshev polynomials. The function should take as inputs, an anonymous function `f` (to be approximated), the degree of approximation `n` and the domain vector `x`.

<!-- grade_id: cell-e4ca5b4d33af2efd, points: 0, grade: false, solution: true -->
```python
# Question 2
def chebApprox(n, x, f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b7674d918052a70d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
f  = lambda x: np.sin(4*x)
n  = 5
x  = np.arange(-1, 1.1, 0.1)
t1 = np.array([ 0.728256,  0.471712,  0.066241, -0.35664 , -0.705994, -0.926386,
               -0.992816, -0.905643, -0.685516, -0.368298, -0.      ,  0.368298,
                0.685516,  0.905643,  0.992816,  0.926386,  0.705994,  0.35664 ,
               -0.066241, -0.471712, -0.728256])
nt.assert_array_almost_equal(t1, chebApprox(5, x, f))
print('Test case passed!!!')
```

<!-- grade_id: cell-b72f8afa93322e2c, points: 9, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 9 marks
```

## Question 3 
Write a function which implements Gaussian Quadrature. Your function should take as inputs, n the order of the Legendre polynomial to use, and f the integrand to be integrated. You may assume that the integrand f has already been transformed to the standard interval.


<!-- grade_id: cell-306d65709af581eb, points: 0, grade: false, solution: true -->
```python
def gaussianQuadrature(n, f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-576ba910561effa6, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# Unit test
f     = lambda t: (4*t + 4)*np.exp(4*t + 4)
ans   = gaussianQuadrature(2, f)
t1    = 3477.5439362670827
nt.assert_almost_equal(ans, t1)
print('Test case passed!!!')
```

<!-- grade_id: cell-ed7cf80d3410d493, points: 13, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 13 marks
```